```python
import os
from numpy.testing import assert_equal
```

# Linux Shell

**NOTE:** For all items, the answer should be a single-line Linux (bash) shell command.  
Concatenation using `&&` and `;` is not allowed.

For this exercise, you will need the demo files.  
Due to the file size limits on ALICE, even under compression, I had to split them into 4 collections: `linux_files.tar.gz.part-aa`, and `linux_files.tar.gz.part-ab`. 

To extract the files, use the command below (either in Jupyter, or in the command line).  
The first command just concatenates the file to remake the full archive.  
The actual extraction is done by the `unzip` command.  
You can remove the gz files once you see this.  

```python
#!cat linux_files.tar.gz.part-* > linux_files.tar.gz && tar -xvf linux_files.tar.gz 
```

## Problem 1
Create the directory structure shown below from the current directory:

    a/b/c
    d/e/f
    j/k/l

```python
%%bash
rm -rf a
rm -rf d
rm -rf j
```

<!-- grade_id: cell-6e94fd9f4fea63f8, points: 0, grade: false, solution: true -->
```python
%%bash
### YOUR CODE HERE

```

<!-- grade_id: cell-891a7af422784155, points: 1, grade: true, solution: false -->
```python
assert_equal(os.path.exists('a/b/c'), True)
assert_equal(os.path.exists('d/e/f'), True)
assert_equal(os.path.exists('j/k/l'), True)
```

## Problem 2
Create a clone copy of the `Sushi Preference Datasets` public data directory to the current directory.  
Being a clone copy, it must preserve all attributes (mode, ownership, timestamps), links and directory structure.

```python
%%bash
rm -rf sushipref
```

<!-- grade_id: cell-f95720c8799e461c, points: 0, grade: false, solution: true -->
```python
%%bash
### YOUR CODE HERE

```

```python
%%bash --bg --out output
ls -ogR --color=never sushipref
```

<!-- grade_id: cell-5dcd5564c6dde16a, points: 1, grade: true, solution: false -->
```python
output = str(await output.read(), encoding='utf-8').strip()
lines_sushi = output.split('\n')
assert_equal(len(lines_sushi), 14)
assert_equal(
    lines_sushi[:5], 
    ['total 440',
     'drwxr-x---@ 10      320 Aug 20 18:31 sushi3-2016',
     '-rw-r--r--   1   221754 Aug 20 14:07 sushi3-2016.zip',
     '',
     'sushipref/sushi3-2016:']
)
!rm -rf sushipref
```

## Problem 3

Rename the directory `y` under the directory `x` to `foo`.

```python
%%bash
rm -rf x
mkdir x
mkdir x/y
mkdir x/z
touch x/y/foo.txt
touch x/z/bar.txt
```

<!-- grade_id: cell-7222f2d9353815aa, points: 0, grade: false, solution: true -->
```python
%%bash
### YOUR CODE HERE

```

```python
%%bash --bg --out output
ls -R --color=never x
```

<!-- grade_id: cell-1339e137e105f4ae, points: 1, grade: true, solution: false -->
```python
output = str(await output.read(), encoding='utf-8').strip()
lines_foo = output.split('\n')
assert_equal(
    lines_foo,
    ['foo', 'z', '', 'x/foo:', 'foo.txt', '', 'x/z:', 'bar.txt']
)
!rm -rf x
```

## Problem 4

Create a symbolic link to `./linux_files/data/public/football-events` on the current directory but name it `soccer`.

```python
%%bash
rm -rf soccer
```

<!-- grade_id: cell-c0233b76c55a6f74, points: 0, grade: false, solution: true -->
```python
%%bash
### YOUR CODE HERE

```

<!-- grade_id: cell-f34d88531e36b2f3, points: 1, grade: true, solution: false -->
```python
assert_equal(os.readlink('soccer'), '/mnt/data/public/football-events/')
```

## Problem 5

Search all files under `./linux_files/data/public/opendata/transport` for lines that contain the phrase `Mactan (International)` (case-insensitive) then store them in the file `mactan.txt`.

```python
!rm -f mactan.txt
```

<!-- grade_id: cell-21ad2d469f3bbbf1, points: 0, grade: false, solution: true -->
```python
%%bash
### YOUR CODE HERE

```

<!-- grade_id: cell-9921b5028941bdc2, points: 1, grade: true, solution: false -->
```python
with open('mactan.txt') as f:
    lines_mactan = sorted(f.readlines())
assert_equal(len(lines_mactan), 83)
assert_equal(
    lines_mactan[:5],
    ['VII (Central Visayas),Mactan (International),Cathay,-,10660,12105,12598,'
     '13782,12508,12844,13541,10803,10921,8329,10486,128577\n',
     'VII (Central Visayas),Mactan (International),Cathay,-,196276,619180,'
     '586180,638074,490074,520351,521658,329835,572562,581802,386086,5442078\n',
     'VII (Central Visayas),Mactan (International),Cathay,-,56,62,46,62,60,62,'
     '62,60,62,60,62,654\n',
     'VII (Central Visayas),Mactan (International),Cathay,11168,10489,-,-,-,-,'
     '-,-,-,-,-,-,21657\n',
     'VII (Central Visayas),Mactan (International),Cathay,299988,420184,449741,'
     '470096,555647,526221,-,466785,486125,464619,403014,-,4542420\n']
)
```

## Problem 6

Update `a.txt` so that the contents of `b.txt` and `c.txt` are appended at the end.

```python
with open('a.txt', 'w') as f:
    f.writelines([
        'little carabao\n',
        'sitting on a chair\n',
        'then she cried aloud\n'
    ])
with open('b.txt', 'w') as f:
    f.writelines([
        'unganga\n',
        'unganga\n',
        'little carabao\n'
    ])
with open('c.txt', 'w') as f:
    f.writelines([
        'three blind mice\n',
        'three blind mice\n',
        'three blind mice\n'
    ])
```

<!-- grade_id: cell-0cddc9937348e941, points: 0, grade: false, solution: true -->
```python
%%bash
### YOUR CODE HERE

```

<!-- grade_id: cell-3dbf03ef87bbec63, points: 1, grade: true, solution: false -->
```python
with open('a.txt') as f:
    lines_abc = f.readlines()
assert_equal(
    lines_abc[:5],
    ['little carabao\n',
     'sitting on a chair\n',
     'then she cried aloud\n',
     'unganga\n',
     'unganga\n']
)
```

## Problem 7

Show the lines in `./linux_files/data/public/gutenberg/7/74/74-0.txt` that contain both case-insensitive strings `river` and `raft`.

<!-- grade_id: cell-75d732e00fa39acf, points: 0, grade: false, solution: true -->
```python
%%bash --bg --out output
### YOUR CODE HERE

```

<!-- grade_id: cell-a96ada60b6ebe6e6, points: 1, grade: true, solution: false -->
```python
output = str(await output.read(), encoding='utf-8')
lines_gb2 = output.strip().split('\r\n')
assert_equal(len(lines_gb2), 3)
assert_equal(
    lines_gb2[0], 
    'places that were in harmony with his spirit. A log raft in the river'
)
```

```python

```